# Dados do aluno
Preencha seus dados nas variáveis da célula abaixo para continuar.

```python
# Preencha seu nome e código do aluno aqui
nome_aluno   = ""   # Substitua aqui pelo seu nome, exemplo:   nome_aluno = "Edson Crosué das Dores"
codigo_aluno = None # Substitua aqui pelo seu código, exemplo: codigo_aluno = 1234567
```

<!-- grade_id: cell-5ef577dcc9470d8c, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert nome_aluno != "" # Você precisa preencher seu nome no formulário
assert codigo_aluno is not None # Você precisa preencher seu código de aluno
```

# Instruções
Neste notebook interativo, você irá aplicar os conceitos aprendidos da aula prática para manipular imagens usando Numpy e OpenCV. 




# Exercício 1: Carregar o módulo cv2 e a imagem
Para manipular imagens, é necessário primeiro carregá-las. Para carregar uma imagem, é necessário que você carregue o módulo `cv2` e, dele, usar a função `imread()`.

No diretório deste caderno, você deve encontrar o arquivo `img_teste.jpg`. 

**Desta forma, carregue a imagem `img_teste.jpg` usando o módulo cv2 em escala de cinza.** Armazene o resultado na variável global `img_teste`.

<!-- grade_id: cell-69594658e5631105, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### (Opcional) Visualize a imagem!
O módulo `matplotlib.pyplot` permite visualizar gráficos em geral. Normalmente "apelidado" de `plt`, para visualizar imagens, o `pyplot` possui a função `imshow()`.

**Utilizando o `pyplot` e o `imshow`, visualize a imagem carregada na célula abaixo.** Lembre-se que um dos parâmetros opcionais é o mapa de cores. Consulte o material de aula para alterar o mapa de cores para cinza.

```python
# Opcional: coloque aqui o código para visualização
```

<!-- grade_id: cell-b742da505d545dc2, points: 3, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
```

# Exercício 2: Região de Interesse
A imagem de entrada possui alguns objetos que estão em cima da minha escrivaninha. Sabendo que o Mega-Man está localizado entre os pontos $(170,315)$ e $(262,450)$, **crie uma nova imagem recortada do Megaman a partir da imagem em `img_teste` e armazene em `img_sol_1`**. 

Chamamos a região de intereste (do recorte) de ROI (do inglês, *Region of Interest*).

Considere que as linhas 170 e 262, e as colunas 315 e 450 estão inclusas na ROI.

<!-- grade_id: cell-c733fdf92389e6e5, points: 0, grade: false, solution: true -->
```python
# Não alterar: célula reservada ao autograder
# YOUR CODE HERE
raise NotImplementedError()
```

### (Opcional) Visualize o resultado!
Como resultado deste exercício, somente o Mega-Man deve aparecer na imagem.

```python
# Opcional: coloque aqui o código para visualização
```

<!-- grade_id: cell-0213d2e97401cd8e, points: 3, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
```

# Exercício 3: Quantização
O bisão Ferdinando, que aparece à esquerda na foto, acha que a imagem ficaria mais estilosa se estivesse quantizada. Faça um favor para ele e **implemente uma função que, dada uma imagem, retorna a mesma quantizada em 4 níveis de cor.** 

Considere que cada intervalo possui 64 níveis de cinza da imagem original.

<!-- grade_id: cell-fb29c3851da405ef, points: 0, grade: false, solution: true -->
```python
def quantiza(imagem) : 
    # YOUR CODE HERE
    raise NotImplementedError()
```

### (Opcional) Visualize o resultado!
A imagem resultante deve possuir somente 5 níveis de cor. Confira se o resultado possui menos cores que a imagem original.

```python
# Opcional: coloque aqui o código para visualização
```

<!-- grade_id: cell-dcb34e627758067a, points: 4, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
```