# Sobre
Este notebook visa apresentar o sistema de avaliações da parte prática da disciplina.

# Identificação
Em todos os notebooks, você terá de se identificar antes de prosseguir para as próximas questões. Altere o valor das variáveis `nome_aluno` e `codigo_aluno` para seu nome (string) e seu código (int).

```python
# Preencha seu nome e código do aluno aqui
nome_aluno   = ""   # Substitua aqui pelo seu nome, exemplo:   nome_aluno = "Edson Crosué das Dores"
codigo_aluno = None # Substitua aqui pelo seu código, exemplo: codigo_aluno = 1234567
```

Caso não preencha esses dados, a célula seguinte irá executar com erros.

<!-- grade_id: cell-93a7b68b991c46ce, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert nome_aluno != "" # Você precisa preencher seu nome no formulário
assert codigo_aluno is not None # Você precisa preencher seu código de aluno
```

A célula acima é uma célula de teste do `autograder`, o sistema que avalia se as respostas estão de acordo com o esperado, e não é possível alterá-las.

Essas células possuem dois tipos de testes: os testes públicos e os testes ocultos. Os **testes públicos** são visíveis para os alunos e avaliam de forma superficial se a resposta está correta.

Para que não seja possível burlar os testes públicos - já que eles costumam dar algum indício do que deve ser a resposta -, existem os **testes privados**, os quais só podem ser vistos pelo professor.

Os alunos são avaliados pelos dois testes, mas no notebook disponível para os alunos só é possível saber se foram aprovados nos testes públicos. Mas fique tranquilo: os testes privados foram feitos para só funcionarem se a implementação para os testes públicos estiver correta.

# Exemplo de Exercício
Existem dois tipos de exercícios: os que as respostas são armazenadas em variáveis e os que as respostas são armazenadas em funções. Em todos os casos, é **evidenciado** no exercício qual é o tipo de resposta esperada (função ou variável).

Vamos implementar um exercício teste em que a **resposta é armazenada em variável**. 

**Rode a célula abaixo. Em seguida, calcule o triplo do valor da variável `exemplo_var` e armazene na variável `resultado_1`**

```python
import random
exemplo_var = random.randint(-2,2)
print(f"O valor em `exemplo_var` é {exemplo_var}")
```

<!-- grade_id: cell-955b630fef635663, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f24af40e627eca4d, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert 'resultado_1' in globals() # Feedback: a variável `resultado_1` não foi declarada!
assert resultado_1 % 3 == 0 # Feedback: o resultado não é múltiplo de 3. Você se lembrou de multiplicar por 3?
```

Perceba os testes públicos avaliaram se a variável foi declarada ou se ela é múltiplo de 3, mas em nenhum momento a resposta verdadeira foi fornecida. Este é o papel dos testes ocultos. 

No teste oculto, um gabarito é fornecida pelo professor e a resposta do aluno é comparada considerando um erro de arredondamento. No caso, o teste oculto desta célula seria somente
```python
gabarito = exemplo_var * 3
assert resultado_1 == gabarito
```

# Outro Exemplo de Exercício
Na maior parte das vezes, os alunos serão responsáveis por implementar funções com a solução, ao invés de variáveis.

Implemente uma função chamada `eleva` que recebe dois números, por exemplo `a` e `b`, e calcule $a^b$. Retorne o resultado.

**Note:** o nome da função **PRECISA** ser `eleva`. Os nomes dos parâmetros, por outro lado, não precisam ser `a` e `b` (a não ser que seja pedido explicitamente).

Em alguns casos a estrutura básica da função pode vir implementada. Em outros casos, você deverá implementar a função do zero.

<!-- grade_id: cell-5d4da13534c3b366, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# (Opcional) Esta célula pode ser usada para testar sua implementação
# Altere `num_1` e `num_2` para verificar o resultado.
num_1 = 2
num_2 = 3
print(f"{num_1} elevado a {num_2} resulta em {eleva(num_1, num_2)}")
```

<!-- grade_id: cell-429589f62d26931b, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert eleva(2,3) == 8 # Feedback: resultado diferente do esperado!
assert eleva(3,2) == 9 # Feedback: resultado diferente do esperado!
assert eleva(3,3) == 27 # Feedback: resultado diferente do esperado!
```

Note que os testes públicos mostram alguns exemplos de respostas esperadas. Nos testes ocultos mais testes são feitos, normalmente aplicando números gerados aleatoriamente e comparando o resultado com uma função-gabarito.

# Validação
No topo do caderno existe um botão escrito `Validate`. Assim que concluir as implementações, rode este botão e ele irá fornecer quais exercícios foram reprovados pelos testes públicos, assim você poderá repensar eles e refazer.

Em alguns casos, pode ser que o botão não apareça. Se for o seu caso, não se preocupe: se algum dos resultados estiver errado de acordo com os testes públicos um erro **em letras garrafais** é lançado logo no teste, então recomendo que reinicie o Kernel e rode novamente todas as células em sequência (dica: clique em `Kernel` e em `Restart & Run All`).